## Problem 1 - Simple temperature calculator (*3 points*)

In the first problem your aim is to create a function that converts the input temperature from degrees Fahrenheit to degrees Celsius. The conversion formula from Fahrenheit to Celsius can be found below.

\begin{equation}
  \Large
  T_{\mathrm{Celsius}} = \left( T_{\mathrm{Fahrenheit}} - 32 \right) / 1.8
\end{equation}

If everything in your script is working properly the following test case should work:

  ```python
  >>> print("32 degrees Fahrenheit in Celsius is:", fahr_to_celsius(32))
  32 degrees Fahrenheit in Celsius is: 0.0
  ```

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable, function, and parameter names as stated in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions. Also, please do not copy/paste existing code cells. You can add cells if you like, but copying and pasting can also cause the automatic tests to fail.

**Your score on this problem will be based on following criteria:**

- Creating a function called `fahr_to_celsius`
- Defining the function to have one input parameter called `temp_fahrenheit`
- Creating a variable called `converted_temp` inside the function to which you should assign the conversion result (i.e., the input Fahrenheit temperature converted to Celsius)
- Returning the converted value from the function back to the user
- Answering some questions about functions at the end of this problem
- Adding comments in your code and a docstring that explains how to use your `fahr_to_celsius` function (i.e., you should write the purpose of the function, parameters, and returned values)
- Uploading your notebook to your GitHub repository for this week's exercise

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Use** your newly created function to print answers to the following questions:

```python
# 1. What is 48째 Fahrenheit in Celsius? 

# 2. What about 71째 Fahrenheit in Celsius?

```

### Problem 1 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells. If you receive an error, you need to change something in the code cells above.

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test1, points: 1, grade: true, solution: false -->
```python
import inspect

# Check that function exists
assert inspect.isfunction(fahr_to_celsius), 'Fahr_to_celsius should be a function.'
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the parameter name is correct
params = list(inspect.signature(fahr_to_celsius).parameters.keys())
assert len(params) == 1, 'The function should have one parameter'
assert params[0] == 'temp_fahrenheit', 'The parameter name should be "temp_fahrenheit".'
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test3, points: 0.5, grade: true, solution: false -->
```python
# Check that the function produces correct answers for:
# 1. What is 48째 Fahrenheit in Celsius? 
assert round(fahr_to_celsius(48), 2) == 8.89
```

<!-- grade_id: problem_1_ag_p1_fahr_to_celsius_test4, points: 0.5, grade: true, solution: false -->
```python
# 2. What about 71째 Fahrenheit in Celsius?
assert round(fahr_to_celsius(71), 2) == 21.67
```

### Check your code

- Make sure you used the given variable names
- Check that you have added necessary comments to your code
- Check that your function has a docstring that describes what it does

### Questions

We would like you to think about and answer the following questions based on the materials and ideas that you learned during the lecture:

  1. Is the concept of function clear to you? If not, what do you not understand?
  2. What are some of the benefits of using functions?
  
Write your answers below (**activate the cell by double clicking**):

<!-- grade_id: problem_1_mg_p2_answers, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 2](Exercise-4-problem-2.ipynb).