## Problem 2 - Temperature classifier (*3 points*)

In [Exercise-3 Problem 2 last week](https://github.com/Geo-Python-2019/Exercise-3/blob/master/Exercise-3-problem-2.ipynb) we practiced how to classify temperatures into four different classes (*i.e. cold, slippery, comfortable, warm*). Let's continue working with the same idea. 

Here, you should create a function called `temp_classifier` that accepts a temperature value in Celsius that will be reclassified into integer numbers 0-3 based on following criteria:

 | Return value | Classification criteria                  |
 |---|-----------------------------------------------------|
 | 0 | temperatures below -2 degrees (Celsius)             |
 | 1 | temperatures equal or warmer than -2, but less than +2 degrees (Celsius)  |
 | 2 | temperatures equal or warmer than +2, but less than +15 degrees (Celsius)  |
 | 3 | temperatures equal or warmer than +15 degrees (Celsius)            |


This problem includes the following steps: 

  - create a new function called `temp_classifier` **(0.5 points)**
  - `temp_classifier` should have one parameter called `temp_celsius` **(0.3 points)**
  - your function should reclassify the input temperature based on the criteria on the above table and return the reclassified value as an integer number (i.e. number 0, 1, 2, or 3) **(1.2 points)**
  - comment your code and add a docstring that explains how to use your `temp_classifier` function (i.e. you should write the purpose of the function, parameters, and return value(s)) **(1.0 points)**
  
    

<!-- grade_id: temp_classifier, points: 0, grade: false, solution: true -->
```python
# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

**Use** your newly created function to print the answers to the following questions:

```python
# 1. What is the class value for 16.5 degrees (Celsius)?

# 2. What is the class value for +2 degrees (Celsius)?
```

### Problem 2 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_2_temp_classifier_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Check that function exists
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_2_temp_classifier_test2, points: 0.3, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the pamameter name is correct
params = list(inspect.signature(temp_classifier).parameters.keys())
assert len(params) == 1
assert params[0] == 'temp_celsius'
```

Check that the function produces correct answers for selected values:

<!-- grade_id: problem_2_temp_classifier_test3, points: 0.3, grade: true, solution: false -->
```python
# 1. What is the class value for 16.5 degrees (Celsius)?
assert temp_classifier(16.5) == 3, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_temp_classifier_test4, points: 0.3, grade: true, solution: false -->
```python
# 2. What is the class value for +2 degrees (Celsius)?
assert temp_classifier(2) == 2, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_temp_classifier_test5, points: 0.3, grade: true, solution: false -->
```python
# 3. What is the class value for +1 degrees (Celsius)?
assert temp_classifier(1) == 1, 'Wrong class'
print("ok :)")
```

<!-- grade_id: problem_2_temp_classifier_test6, points: 0.3, grade: true, solution: false -->
```python
# 4. What is the class value for -5 degrees (Celsius)?
assert temp_classifier(-5) == 0, 'Wrong class'
print("ok :)")
```

Did you check that:
 - all the variable names are correct?
 - your code has enough comments?
 - your function has an informative dosctring?

<!-- grade_id: problem_2_manual_check, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 3](Exercise-4-problem-3.ipynb).